```python
# Initialize Otter
import otter
grader = otter.Notebook("hw01.ipynb")
```

# Toward functional python
In this worksheet, we'll practice with some more advanced forms of Python, including list and set comprehensions, anonymous functions, mapping, and filtering. These are all necessary to accomplish "functional programming" in Python. Note that the tests in this section only test that your code works. For full credit, code must also comply with the limits listed for each question. These limits get more stringent as the questions continue. 

<!-- BEGIN QUESTION -->

Question 1: in the cell below, write a function that takes a list of numbers as an argument and filters the list so that numbers more than two standard deviations from the mean on either side are omitted. Use the large sample standard deviation, to wit: $\sqrt((\sum_1^n (\mu - x) ^ 2))/n)$ where $n$ is the number of list elements and $\mu$ represents their mean. For this question, you can use any method to compute the result. You are allowed to start by using `for` loops for each part, using `list.append` to build the result, the way you would in CS 11 or CS 15.

```python
from math import sqrt
def clean1(data):
    result = []
    length = len(data)
    mean = 0
    dev = 0
    for x in data:
        mean += x
    mean = mean/length
    for x in data:
        dev += (mean - x) ** 2
    dev = sqrt(dev/length)
    for x in data:
        if x < (mean + (2 * dev)) and x > (mean - (2 * dev)):
            result.append(x)
    return result
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 2: Repeat question 1, but eliminate `for` loops in favor of list comprehensions and the use of `sum` for sums. 

```python
def clean2(data):
    length = len(data)
    mean = sum(data)/length
    d = [(mean - x)**2 for x in data]
    dev = sqrt(sum(d)/length)
    result = [x for x in data if x < (mean + (2 * dev)) and x  > (mean - (2 * dev))]
    return result
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 3: Repeat question 1, but this time, eliminate list comprehensions *and* `for` loops in favor of the "functional programming" functions map, filter, and lambda. 

```python
def clean3(data):
    length = len(data)
    mean = sum(data)/length
    d = map(lambda x: (mean - x) ** 2, data)
    dev = sqrt(sum(d)/length)
    result = filter(lambda x: x < (mean + (2 * dev)) and x > (mean - (2 * dev)), data)
    return list(result)
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 4: Test whether your solution for Problem 3 works on a `set`, and indicate results below. What works? What doesn't?

_Type your answer here, replacing this text._

```python
# Case 1: List and Set contains NO DUPLICATES and INCREASING ORDER
# Observations: 
# 1. Same output numbers 
# 2. Numbers are printed based on the order of the list/set given (left to right)
print("Case 1 No duplicates + Increasing order: 1, 4, 8, 10, 21, 56")
print("List: " + str(clean3([1,4,8,10,17,56])))
print("Set: " + str(clean3({1,4,8,10,17,56})) + "\n")

# Case 2: List and Set contains NO DUPLICATES and RANDOM ORDER
# Observations: 
# 1. Same output numbers 
# 2. Orders are different. List output follows the order of the list given. Set output order was sorted in increasing order
print("Case 2 No duplicates + Random order: 1, 10, 4, 8, 17, 56")
print("List: " + str(clean3([1, 10, 4, 8, 17, 56])))
print("Set: " + str(clean3({1, 10, 4, 8, 17, 56})) + "\n")

# Case 3: List and Set contains DUPLICATES and INCREASING ORDER
# Observations: 
# 1. Different output numbers - List has 56, Set does not. List contains duplicates, Set does not
# 2. Numbers are printed based on the order of the list/set given (left to right)
print("Case 3 Yes duplicates + Increasing order: 1, 4, 4, 8, 10, 17, 56, 56")
print("Duplicate List: " + str(clean3([1, 4, 4, 8, 10, 17, 56, 56])))
print("Duplicate Set: " + str(clean3({1, 4, 4, 8, 10, 17, 56, 56})) + "\n")

# Case 4: List and Set contains DUPLICATES and RANDOM ORDER
# Observations: 
# 1. Different output numbers - List has 56, Set does not. List contains duplicates, Set does not
# 2. Orders are different. List output follows the order of the list given. Set output order was sorted in increasing order
print("Case 4 Yes duplicates + Random order: 4, 1, 4, 56, 8, 17, 10, 56")
print("Duplicate List: " + str(clean3([4, 1, 4, 56, 8, 17, 10, 56])))
print("Duplicate Set: " + str(clean3({4, 1, 4, 56, 8, 17, 10, 56})) + "\n")

print("WHAT WORKS")
print("- A set correctly calculates mean and standard deviations for non duplicate containing sets")
print("- A set will return a sorted output of values in increasing order \n")
print("WHAT DOESN'T WORK")
print("- A set does not correctly calculate mean and standard deviations for duplicate containing sets. ")
print("- This is because sets do not hold duplicate items, so a number will only count once and another of the same number will be disregarded. ")
print("- This will then affect the calculation of the mean and standard deviation. In these test cases, a duplicate set from case 3&4 will have the same output as case 1&2. \n")
print("SUMMARY")
print("Non duplicate set = correct output\nDuplicate list = correct output, Duplicate set = wrong output\nDuplicate set output = Non duplicate set output\nAny set = sorted output")

```

Output:
```
Case 1 No duplicates + Increasing order: 1, 4, 8, 10, 21, 56
List: [1, 4, 8, 10, 17]
Set: [1, 4, 8, 10, 17]

Case 2 No duplicates + Random order: 1, 10, 4, 8, 17, 56
List: [1, 10, 4, 8, 17]
Set: [1, 4, 8, 10, 17]

Case 3 Yes duplicates + Increasing order: 1, 4, 4, 8, 10, 17, 56, 56
Duplicate List: [1, 4, 4, 8, 10, 17, 56, 56]
Duplicate Set: [1, 4, 8, 10, 17]

Case 4 Yes duplicates + Random order: 4, 1, 4, 56, 8, 17, 10, 56
Duplicate List: [4, 1, 4, 56, 8, 17, 10, 56]
Duplicate Set: [1, 4, 8, 10, 17]

WHAT WORKS
- A set correctly calculates mean and standard deviations for non duplicate containing sets
- A set will return a sorted output of values in increasing order 

WHAT DOESN'T WORK
- A set does not correctly calculate mean and standard deviations for duplicate containing sets. 
- This is because sets do not hold duplicate items, so a number will only count once and another of the same number will be disregarded. 
- This will then affect the calculation of the mean and standard devi
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

You are not done until you upload the exported zipfile to GradeScope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1 (points: 2)
```
>>> assert clean1([5,6,7,6,5,20]) == [5,6,7,6,5]
```
```
>>> assert clean1([5,20,6,7,6,5]) == [5,6,7,6,5]
```
```
>>> assert clean1([20,5,6,7,6,5]) == [5,6,7,6,5]
```


Tests for q2 (points: 1)
```
>>> assert clean2([5,6,7,6,5,20]) == [5,6,7,6,5]
```
```
>>> assert clean2([5,20,6,7,6,5]) == [5,6,7,6,5]
```
```
>>> assert clean2([20,5,6,7,6,5]) == [5,6,7,6,5]
```


Tests for q3 (points: 1)
```
>>> assert clean3([5,6,7,6,5,20]) == [5,6,7,6,5]
```
```
>>> assert clean3([5,20,6,7,6,5]) == [5,6,7,6,5]
```
```
>>> assert clean3([20,5,6,7,6,5]) == [5,6,7,6,5]
```
