```python
# Initialize Otter
import otter
grader = otter.Notebook("hw02.ipynb")
```

# Structures and dictionaries
In this homework, we'll explore structures and in particular dictionaries, including what they might mean and what we can do with them. We start with a Facebook-like "likes" dictionary of who likes who. This is not reflexive. The fact that X likes Y does not mean that Y likes X! Run the following cell to create such a dictionary. 

```python
likes = { 
     'Amy': {'Mike'}, 
     'John': {'Mike', 'Amy'}, 
     'Emily': {'Amy'},
     'George': {'Frank', 'Emily', 'Mae'}, 
     'Mae': {'John'}
}
from pprint import pprint
pprint(likes)
```

Output:
```
{'Amy': {'Mike'},
 'Emily': {'Amy'},
 'George': {'Mae', 'Frank', 'Emily'},
 'John': {'Mike', 'Amy'},
 'Mae': {'John'}}
```

<!-- BEGIN QUESTION -->

Question 1: in the cell below, write a function that takes one argument that is a dictionary structured like the one above, and returns the whole set of people involved in the dictionary as either keys or parts of values. 

```python
def everyone(data):
    result = set()
    for x in data:
        result.add(x)
        for y in data[x]:
            result.add(y)
    return result
```

```python
everyone(likes)
```

Output:
```
{'Amy', 'Emily', 'Frank', 'George', 'John', 'Mae', 'Mike'}
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 2:  Write a function to determine who likes a specific person, based upon data in the dictionary. This is a matter of inverting the dictionary. 

```python
def wholikes(person, data):
    result = set()
    for x in data:
        for y in data[x]:
            if person == y:
                result.add(x)
    return result
```

```python
# Use this to test your code
print("who likes John?")
pprint(wholikes('John', likes))
print("who likes Amy?")
pprint(wholikes('Amy', likes))
```

Output:
```
who likes John?
{'Mae'}
who likes Amy?
{'Emily', 'John'}
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 3: Write a function to invert the whole structure so that it describes who is liked rather than who likes who. In this structure -- call it `d` -- `d[person]` is the number of people who like `person`. Hint: if there's not a set for a person, you have to create it. Test whether there's a set via `person in output`.

```python
def invert(data):
    result = dict()
    for x in data:
        for y in data[x]:
            if y in result:
                value = result.get(y)
                value.add(x)
                result.update({y: value})
            else:
                result.update({y: {x}})
    return result
```

```python
# Use this cell to test your code: 
pprint(invert(likes))
pprint(invert(dict()))
```

Output:
```
{'Amy': {'Emily', 'John'},
 'Emily': {'George'},
 'Frank': {'George'},
 'John': {'Mae'},
 'Mae': {'George'},
 'Mike': {'John', 'Amy'}}
{}
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

Question 4: Does it make any difference at all whether we use sets or lists inside `likes`? What is the difference if we use lists? 

_Type your answer here, replacing this text._

```python
'''It makes no difference if we use sets or lists as long as there are no duplicates; Meaning, person A cannot like person B twice. 
 In a case where duplicates occur (which shouldn't), we would get duplicate results; 
 For example, if A likes B and B, the invert function will result in B being liked by A and A
  In summary, there is no difference as long as the likes follow the logic of being able to like a person only once.
'''
```

Output:
```
"It makes no difference if we use sets or lists as long as there are no duplicates; Meaning, person A cannot like person B twice. \n In a case where duplicates occur (which shouldn't), we would get duplicate results; \n For example, if A likes B and B, the invert function will result in B being liked by A and A\n  In summary, there is no difference as long as the likes follow the logic of being able to like a person only once.\n"
```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

You are not done until you upload the exported zipfile to GradeScope.

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1 (points: 1)
```
>>> assert everyone(likes) == {'Amy', 'Mike', 'John', 'Amy', 'Emily', 'George',  'Frank', 'Mae'}
```
```
>>> assert everyone({'foo': {'bar'}}) == {'foo', 'bar'}
```
```
>>> assert everyone({'foo': {}}) == {'foo'}
```
```
>>> assert everyone(dict()) == set()
```


Tests for q2 (points: 1)
```
>>> assert wholikes('Mike', likes) == {'Amy', 'John' }
```
```
>>> assert wholikes('cat', {'dog': {'cat'}}) == {'dog'}
```
```
>>> assert wholikes('dog', {'dog': {'cat'}}) == set()
```
```
>>> assert wholikes('dog', dict()) == set()
```


Tests for q3 (points: 2)
```
>>> assert invert({'dog': {'cat', 'rat'}}) == {'cat': {'dog'}, 'rat':{'dog'}}
```
```
>>> assert invert({'dog' : {'cat'}, 'rat': {'cat'}}) == {'cat' : {'rat', 'dog'}}
```
```
>>> assert invert(dict()) == dict()
```
